```python
# Initialize Otter
import otter
grader = otter.Notebook("hw06.ipynb")
```

# Homework 6: Probability Distributions



## Due Date: Wednesday, Feb 18th by 11:59 PM MT on Gradescope

## Assignment Learning Objectives
 - Explain the difference between Probabilty Density Functions (PDFs) and Cumulative Density Functions (CDFs) and use both to calculate probabilities for continuous random variables
 - State the Probability Mass Functions, Expected Value and Variance for Bernoulli, Binomial and Poisson RV and use them to calculate probabilities
 - Distinguish when to use Bernoulli vs Binomial vs Poission Random Variables to model a given situation and state any assumptions that are needed to use these RV
 - Use Python to simulate sampling from a Bernoulli, Binomial or Poisson Distribution
 - State the PDF, Expected Value and Variance for a Uniform continuous random variable and use it to calculate probabilities
 - State the PDF, Expected Value and Variance for an Exponential continuous random variable and use it to calculate probabilities
 - State the PDF, Expected Value and Variance for Normal continuous random variable and use it to calculate probabilities
 - Distinguish when to use Uniform vs Exponential vs  Normal random variables to model a given situation and state any assumptions that are needed to use these RV



## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Below are examples of acceptable vs unacceptable use of resources and collaboration when doing HW assignments in CSCI 3022.


The following would be some **examples of cheating** when working on HW assignments in CSCI 3022.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given homework problem or one similar is cheating.  However, consulting the class notes, and web pages that explain the material taught in class but do NOT show a solution to the homework problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course notes, the course textbook, and references that contain syntax and/or formulas.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the HW before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by a TA/Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit HW Hours and/or post questions on Piazza!


#### Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


    

### Score breakdown

Question | Points | Grading Type
--- | --- | ---
Question 1adf | 7 | manual
Question 1bce | 4 | auto
Question 2 | 14 | manual
Question 3 | 7 | autograded
Question 4 | 8 | manual
Question 5a | 2 | manual
Question 5b | 2 | auto
Question 6a | 2 | manual
Question 6b | 1 | auto
Question 6c | 3 | manual and auto
|Total | 50



## Assigned Readings for this HW:

- [Random Variables](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/rvs/)
- [Probability Mass Functions](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/pmf/)
- [Expected Value](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/expectation/)
- [Variance](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/variance/)
- [Bernoulli Distribution](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/bernoulli/)
- [Binomial Distribution](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/binomial/)
- [Poisson Distribution](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/poisson/)
- [Continuous Random Variables](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/continuous/)
- [Probability Density](https://data88s.org/textbook/content/Chapter_10/01_Density.html)
- [Continuous Uniform Distribution](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/uniform/)
- [Exponential Distribution](https://www.probabilitycourse.com/chapter4/4_2_2_exponential.php)
- [Normal Distribution](https://onlinestatbook.com/2/normal_distribution/normal_distribution.html)


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

```python
import numpy as np
import pandas as pd
import math
import scipy.stats as stats
from scipy import special
import matplotlib
import matplotlib.pyplot as plt
from matplotlib.ticker import PercentFormatter
import seaborn as sns
import hashlib
from IPython.display import display, Latex, Markdown
import simple_latex_checker as slc

plt.style.use('fivethirtyeight')
# Default plot configurations
plt.rcParams['figure.figsize'] = (16,8)
plt.rcParams['figure.dpi'] = 150
sns.set()
import warnings
warnings.simplefilter('ignore', FutureWarning)


def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()

%matplotlib inline
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴 Question 1 ##

In the general population, approximately $10\%$ of people are left-handed.  You randomly sample (with replacement) $77$ people from the general population.  


### 🔴 Question 1a (3 pts)  ##
Answer each of the following questions using Markdown in the same cell below.



 - **i).**  Let $X$ be the random variable that is the number of people in your sample who are left handed.   Define how $X$ is distributed using the notation: $X\sim DistributionNickName(parameter\_values)$.  

 
 - **ii).** What is the support of $X$?

 -  **iii).**  What is the PMF of $X$? (Give your answer as a closed-form function)


Answer all of the questions above in the same LaTeX cell below. 





_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 🔴 Question 1b (2 pts)  ##
Calculate the following in the code cell below. 

For all parts of this question do NOT use any built-in PMF or CDF methods in Python (i.e. do not use Scipy).   
Also, **do NOT use the math.factorial** method (it will fail in gradescope).  
You can use either `special.comb` or `special.factorial` as needed.  

 - **i).**  What is the probability $9$ people in your sample are left-handed?  



 - **ii).**  What is the probability that **at least** $7$ people and **at most** $10$ people in your sample are left-handed?  

 - **iii).**  What is $E[X]$?   (you can use the formula we derived in class)
 
 - **iv).**  What is $SD[X]$?    (you can use the formula we derived in class)


```python
...

#i). 
prob_9 = ...
print("Prob of exactly 9 left-handed people = ", round(prob_9,3))

#ii).
prob_7_to_10 = ...
print("Prob of between 7 and 10 left-handed people = ", round(prob_7_to_10,3))

#iii).
ev_X = ...
print("Expected Value of X = ", ev_X)

#iv).
sd_X = ...
print("Standard Deviation of X = ", round(sd_X,3))

```

```python
grader.check("1b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 1c (1 pt)

### Built-In Python Functions

Python's scientific computing library `scipy` has a `stats` module with built-in functions for many commonly used Random Variables:
https://docs.scipy.org/doc/scipy/reference/stats.html

You'll see in the first cell of code above we have already imported this module with the alias stats: `import scipy.stats as stats`


 - Read the documentation for `stats.binom`: https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.binom.html.
 
 - Then recalculate part 1b(i) above using `stats.binom.pmf()`.  (That is, calculate the probability exactly $9$ people in your sample are left-handed)

```python
#Recalculate part 1b(i) above using stats.binom.pmf()

bi_recalculated = ...

round(bi_recalculated,3)
```

```python
grader.check("1c")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 1d (1 pt)

To plot a histogram of the distribution of $X$, 
use  `stats.binom.pmf` to define the pmf array in the code below.  
Be sure to also fill in the missing parameters in the title of the plot.

```python

X = np.array(range(78)) 

pmf = ...

fig, ax = plt.subplots()

# Note: to plot this distribution we can't use .hist, because we don't 
# have an array of data, 
# Instead we use .bar because we have an array of probabilities 

ax.bar(X, pmf);

plt.xlim(0,77)

# Fill in the correct values for n and p in the title:
plt.title("Binomial Distribution,n=..., p=...");


```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 1e (1 point)


The Cumulative Distribution Function (CDF), $F(a)$ of a Random Variable $X$ is defined as:

$F(a) = P(X \leq a)$

`scipy.stats` has the CDF as a built-in function.  


**Use `stats.binom.cdf()` to calculate the following probabilities:** 

 - **i).** What is the probability that **at most** $10$ people in your sample are left-handed?

 - **ii).**  Re-do question 1b(ii) using the CDF function:  What is the probability that **at least** $7$ people and **at most** $10$ people in your sample are left-handed?    (Hint: you will need to use subtraction as part of your calculation).


```python
prob_at_most_10 = ...
print(round(prob_at_most_10, 3))

prob_7_to_10 = ...
print(round(prob_7_to_10, 3))
```

```python
grader.check("q1e")
```

**Randomly Sampling From a Distribution** 

`scipy.stats` has built-in methods `.rvs` to generate random samples from common distributions.  

To generate random samples from a Binomial distribution we can use `stats.binom.rvs(n, p, size)` , where size is the size of the random sample (assuming with replacement()


   
 -  Run the code below to randomly sample 15 times from a Binomial distribution with $n=25$ and $p = 0.20$:

```python
x = stats.binom.rvs(n=20, p=0.20, size=15)
x
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 1f (3 points)

We can approximate the theoretical distribution of _Bin(n,p)_ by drawing many random samples from the distribution and plotting a **density** histogram of the sample data. 

In the code below, you are going to repeatedly draw random samples from a Bin(77, 0.10) distribution using sample sizes of 50, 500 and 5000 and plot the resulting empirical distributions of your samples.  

You'll then compare this to the theoretical distribution you plotted in 1d.

Fill in all the ellipses (...) in the code below to complete this.


You should see that the empirical distributions approach the theoretical distribution as the sample size increases.  





```python

#Randomly sample 50 values from Bin(77,0.10)
sample_50 = ...

#Randomly sample 500 values from Bin(77,0.10) 
sample_500 = ...


#Randomly sample 5000 values from Bin(77,0.10) 
sample_5000 = ...

# Create a figure that will have 4 subplots (4 histograms))
fig, ax = plt.subplots(nrows=4, ncols=1, figsize = (3,5))


plt.subplots_adjust(hspace=1)

# Plot a density histogram of the sample data of size 50
# ax[0] indicates we are going to plot this distribution in the first out of 4 subplots
# Use density histogram because we are binning a list of simulated numbers to estimate the probabilities
ax[0].hist(sample_50, density=True, bins = np.arange(-0.5, 100, 1))
ax[0].set_axisbelow(True)
ax[0].grid(alpha=0.25)
ax[0].set_xlim([0, 25]);
ax[0].set_title("Simulated Bin(77, 0.10) distribution using sample of size 50")



# Use the same format as above to plot 
# a density histogram of the sample data of size 500.
# Note, to put this as the 2nd plot out of 4, use ax[1] (instead of ax[0])
...
ax[1].set_axisbelow(True)
ax[1].grid(alpha=0.25)
ax[1].set_xlim([0, 25]);
ax[1].set_title("Simulated Bin(77, 0.10) distribution using sample of size 500")

# Plot a density histogram of the sample data of size 5000
...
ax[2].set_axisbelow(True)
ax[2].grid(alpha=0.25)
ax[2].set_xlim([0, 25]);
ax[2].set_title("Simulated Bin(77, 0.10) distribution using sample of size 5000")


# Plot the actual theoretical distribution.  Use bar because we know the exact probabilities:
...
ax[3].set_axisbelow(True)
ax[3].grid(alpha=0.25);
ax[3].set_xlim([0, 25]);
ax[3].set_title("Theoretical Binomial Distribution, n=77, p=0.10")

plt.show()
```

<!-- END QUESTION -->


You should see that the simulated/empirical distributions approach the theoretical distribution as the sample size increases.  

<img src="img/int.png" width="600"> 

<!-- BEGIN QUESTION -->


<a id='p3'></a>


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴 Question 2 (14 points) ##


Let $X$ be a continuous random variable with support $1 \leq x \leq 7$.  

The graph of the PDF $f(x)$ for $X$ is shown here:

<img src="img/PDF.png" width="500"> 

a).  What is $P(3.25\leq X \leq 5)$?  Show all work using LaTex and give your final answer as a fully simplified fraction.  

b).  What is $E[X]$?  Show all work using LaTeX and give your final answer as a fully simplified fraction.

c).  What is $Var[X]$?  Show all work using LaTeX and give your final answer as a fully simplified fraction.

d).  What is $SD[X]$?  

Note - to receive full credit you must:
     
     1. Show all work used to find the equation for each of the lines in the piecewise function - do not just list the equation of the lines without supporting work.
     
     2.  Show all steps for each integral calculation (see the picture guide above for what counts as full steps).



Answer all 4 questions in the same cell below using LaTeX.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answers to 2 enter it below.  Note that this autograder test is not worth any points, it's just for you to check your work above to make sure you're on the right track. 

```python
answer_2a = ...
```

```python
answer_2b = ...
```

```python
answer_2c = ...
```

```python
answer_2d = ...
```

```python
grader.check("q2_check")
```

[Back to top](#top)

<a id='p2'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴 Question 3: Poisson & Exponential Distributions Video Assignment (7 pts) ## 

#### Complete this [Video Assignment (click this link)](https://canvas.colorado.edu/courses/133218/assignments/2588782) available on Canvas (to receive points you must watch the entire video and answer the embedded questions).

The slides for this video are from Lesson 12 and 14 and are located in the Canvas Modules.
As you watch the video, fill in the code below along with the video so you can practice with the concepts in the video.

<img src="img/Poisson.png">


**Example:**  There are an average of 2.79 major earthquakes in the world each year, and major earthquakes occur independently (assume we're ignoring aftershocks).  What is the probability of exactly 3 major earthquakes happening next year?

```python
# Calculate using the Poisson PMF formula

...
```

We can also calculate this directly using built-in PMF function in Scipy:
https://docs.scipy.org/doc/scipy/reference/stats.html 


```python
# Calculate using Scipy built-in PMF function

...


```

Plot the PMF of $X\sim Poi(3)$ 

```python
...
```

What is the probability that **at least** 3 earthquakes happen next year?  Calculate using the PMF formula.  


```python
...
```

```python
...
```

We can also calculate this directly using built-in CDF function in Scipy:

Use Scipy's CDF to calculate the probability that **at least** 3 earthquakes happen next year.    

```python
...
```

Plot the PMF of $X\sim Poi(\lambda)$, for various different values of $\lambda$, and compare/contrast how $\lambda$ affects the distribution.

```python
...
```

**Randomly sampling from a Poisson distribution:** 

Use Scipy to randomly sample 10 datapoints from a Poisson distribution with $\lambda = 2.79$

```python
...
```

Plot the distribution of a sample of 100 datapoints randomly sampled from a Poisson distribution with $\lambda = 2.79$

```python
...
```

Let's compare the empirical (i.e. simulated) distribution with the theoretical distribution.  Experiment with how the empirical distribution changes as the sample size increases.

```python
#Simulated Values
sample_size = ...
sample = ...
bins = np.arange(-0.5, 15, 1)

#Actual theoretical probabilities
max = 15
X = np.array(range(max+1))
pmf = ...

fig, ax = plt.subplots(nrows=2, ncols=1)
plt.subplots_adjust(hspace=0.5)

# Use histogram because we are binning a list of simulate numbers to calculate the probabilities
ax[0].hist(sample, edgecolor='white', density=True, bins=bins)
ax[0].set_axisbelow(True)
ax[0].grid(alpha=0.25)
ax[0].set_xlim([0, 15]);
ax[0].set_title("Simulated Poisson Distribution: lambda = 2.79")


# Use bar because we know the probabilities
ax[1].bar(X, pmf, edgecolor='white', width=1)
ax[1].set_axisbelow(True)
ax[1].grid(alpha=0.25);
ax[1].set_xlim([0, 15]);
ax[1].set_title("Theoretical Poisson Distribution, lambda = 2.79")

plt.show()
```

## The Exponential Distribution with Python
***

<img src="img/Exponential.png">

Usage note!  Check your documentation!  We have been using the following PDF for the exponential:

$$f(x) = \lambda e^{-\lambda x} \text{  for  } x \in [0, \infty)$$

where $\lambda$ is the **rate** at which events happen.

... but scipy uses a variant that instead has parameter called *scale* that's $1/\lambda$ instead.  See documentation:

https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.expon.html

- A common parameterization for expon is in terms of the rate parameter lambda, such that pdf = lambda * exp(-lambda * x). This parameterization corresponds to using scale = 1 / lambda.

They're in the wrong here, but what can we do?  We have to use `scale=1/lambda` for our usage.


###  Calculating probabilities for an exponential distribution
***  





The scipy stats module has the following built-in functions for the exponential distribution:

 - pdf:  `stats.expon.pdf(x, scale=1/lambda)`

 - cdf:  `stats.expon.cdf(x, scale=1/lambda)`

 - ppf: (for finding percentiles):  `stats.expon.ppf(probability, scale=1/lambda)`

 - rvs:  for randomly sampling from an exponential distribution:  `stats.expon.rvs(samplesize=1000, scale=1/lambda)`


### Exercise 2:
**In an exponential distribution, let $X$= amount of time (in minutes) a postal clerk spends with a customer.   The time is known from historical data to have an average amount of time equal to four minutes, with parameter λ=0.25.    What is the probability that a clerk spends four to five minutes with a randomly selected customer?!**




```python
#Your answer here:
...
```

Let's visualize this using the graph of the pdf and the `plt.plot` function

```python
#plot the exponential pdf
x = np.linspace(0, 10, 1000)

y = stats.expon.pdf(x, scale=4)

plt.plot(x,y )
#shade the region between x=4 and x=5:
plt.fill_between(x, stats.expon.pdf(x, scale=4), where=[(x > 4) and (x < 5) for x in x], label='area is approx 0.08', alpha=0.4)


#choose where to position the legend
plt.legend(loc=(.60,.6), shadow=True)

#add title:
plt.title("PDF of X~exp(0.25)")
```

<!-- BEGIN QUESTION -->

[Back to top](#top)

<a id='p4'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴 Question 4 (8 points)

A factory manufactures solid state hard drives for Seagate Technology. 

For each scenario described below: 
 - i).  State the distribution you could use to model this random variable in the form $RandomVariable\sim DistributionName(parameter\_values)$
 
 - ii).  State any assumptions you need to make to use this distribution to model this situation.  

 - iii). State the support of the Random Variable.



 - iv).  Give the probability distribution for the random variable (i.e. either the pmf or pdf) as a closed-form function.


 

SCENARIOS:


**a).**    Approximately $4\%$ of hard drives that come off of a particular assembly line are defective.  Let $Y$ be the number that are defective in a batch of $300$.  

**b).**  Approximately $4\%$ of hard drives that come off of a particular assembly line are defective.   Let $W$ be whether or not a hard drive is defective.  

**c).**  On average, there are  $8$ defective drives per day.  Let $X$ be the time (in hours) between defective drives.

**d).**  On average, there are $8$ defective drives per day.  Let $M$ be the number of defective drives in a day.  


Answer all parts (a-d) below in the same cell:



_Type your answer here, replacing this text._

<!-- END QUESTION -->

[Back to top](#top)

<a id='p5'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴 Question 5

The Huffmeister floodplane in Houston has historically been estimated to flood at an average rate of $1$ flood every $500$ years. A flood plane with that rate of flooding is called a $500$ year floodplane.      You can assume the time that each flood occurs is independent and the average number of floods over time is constant.  

<!-- BEGIN QUESTION -->

### 🔴 Question 5a (2 points)

What is the probability that at least one flood will occur within the next 90 years?   Fully justify/show all steps and write your answer using LaTeX.  Round your final answer to 5 decimal places.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer to part (a), enter it in the cell below.  Note that this check is not worth any points.  To receive credit you must fully justify your answer in the cell above.

```python
q5a_answer = ...

```

```python
grader.check("q5a_check")
```

### 🔴 Question 5b (2 points)

What is the **expected number** of years until the next flood?

```python
q5b_answer = ...
```

```python
grader.check("q5b")
```

[Back to top](#top)

<a id='p6'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴 Question 6

The weekly amount of money spent on maintainence and repairs by a company was observed over a long period of time to be approximately normally distributed with mean $\$500$ and variance $\$400$.   

Let $W$ represent the weekly amount of money spent on maintenance/repairs.  

That is, $W \sim N(500,400)$.




<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴 Question 6a (2 points)

Graph the pdf of $W$ on the domain $400\leq W \leq600$ 

(Hint: use `plt.plot` and an appropriate method from  Scipy's `stats.norm`: 
https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.norm.html). **Be sure to carefully read what parameters the method takes as inputs.**



```python

# Create a vector of w values between 400 and 600
w = np.linspace(400, 600, 1000) 

...
# Your code above this line


plt.title("PDF of W ~ N(500, 400)")
plt.ylabel("probability density")
plt.xlabel("W")


plt.show()
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴 Question 6b (1 point)

If $\$475$ is budgeted for next week, what is the probability that the actual costs will exceed the budgeted amount?   

Use an appropriate method from   Scipy's `stats.norm` to calculate this:
https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.norm.html



```python
exceed_budget = ...

print("P(W > 475) = {:.4f}".format(exceed_budget))
```

```python
grader.check("q6b")
```

We can show graphically what your answer above means using the following code.  Read through the code so you understand what it is doing.

```python


w = np.linspace(400, 600, 1000)

# Plot the PDF 
plt.plot(w, stats.norm.pdf(w,500,20))

# Shade the area that represents the probability we calculated
(
plt.fill_between(w, stats.norm.pdf(w,500,20), 
                 where=[(w > 475)  for w in w], 
                 label='P(W>475) = '+str(round(exceed_budget,2)), 
                 alpha=0.4)
)


# Add labels to show where the shaded region starts/ends
labels=[400, 475, 600]
plt.xticks(labels)
plt.legend()
plt.title("PDF of W ~ N(500, 400)")
plt.ylabel("probability density")
plt.xlabel("W")




plt.show()
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴 Question 6c (3 points)

i).  What is the probability the actual costs are between  $\$475$ and $\$515$?   Use an appropriate method from Scipy's stats.norm to calculate this: https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.norm.html


ii).  Then show graphically what this number means in terms of the PDF (for full credit your output should have the same format and labels as the plot shown directly above). 

```python
q6ci_answer = ...
```

```python
# Code to create graph for 5c(ii)

...

plt.show()
```

```python
grader.check("q6c")
```

<!-- END QUESTION -->

***

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 6!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any additional websites you referenced*

### Submission Instructions

Before proceeding any further, **save this notebook.**

## IMPORTANT STEP BEFORE COMPILING THIS ASSIGNMENT
To make sure all of your plots render correctly in the PDF, please do the following steps:

1.) In the toolbar, go to Settings  
2.) Disable Autosave Documents

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 


* **Homework 6 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW6 Autograded assignment in Gradescope.

* **Homework 6 Manually Graded**: Submit your hw06.PDF to the HW6 Manually Graded assignment in Gradescope.  **It is your responsibility to fully review your PDF file before submitting and make sure that all your lines of code are visible and any LaTeX has correctly compiled and is fully viewable.**  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.** If not, you will lose points.    

**TROUBLESHOOTING TIPS**  If you are having any issues compiling your assignment, please [read through these troubleshooting tips first](https://docs.google.com/document/d/1ndr3Wj1PSF5qzlLMaBJznwh6QGeEXjd5TAJ6nf9EJvo/edit?usp=sharing), then post any questions on Piazza.

If you are noticing an issue with your plots not printing to the PDF here is the fix:  In the toolbar go to Settings.  Then  Disable autosave.  This is enabled by default in JupyterLab. 


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

```python
nb = slc.Nb_checker()
nb.run_check("hw06.ipynb")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw06.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for 1b (points: 2)
```
>>> assert get_hash(round(prob_9, 3)) == '8eb69fe48ceeb8bb7d08789dabd30eff'
```
```
>>> assert get_hash(round(prob_7_to_10, 3)) == 'f83293212c36437b1f91cca25578929d'
```
```
>>> assert get_hash(round(ev_X, 3)) == '971455e5b307f416baac0cb24b710792'
```
```
>>> assert get_hash(round(sd_X, 3)) == 'a2ca131cd2dad5e1d859fbd3f7d9349d'
```


Tests for 1c (points: 1)
```
>>> assert get_hash(round(bi_recalculated, 3)) == '8eb69fe48ceeb8bb7d08789dabd30eff'
```


Tests for q1e (points: 1)
```
>>> assert get_hash(round(prob_at_most_10, 3)) == '31848e7a15e5da22bfefea472718c52e'
```
```
>>> assert get_hash(round(prob_7_to_10, 3)) == 'f83293212c36437b1f91cca25578929d'
```


Tests for q2_check (points: 0)
```
>>> assert get_hash(round(answer_2a, 6)) == '4b4de284a48f14f2be1c54ce4a62b451'
```
```
>>> assert get_hash(round(answer_2b, 6)) == '60d3837aaa04c5f9aa8bdbe367c34bd7'
```
```
>>> assert get_hash(round(answer_2c, 6)) == '1955257ad9723502516b5cff368dc184'
```
```
>>> assert get_hash(round(answer_2d, 6)) == '0b9cec8fa29d5e16e35fcebf260bba2d'
```


Tests for q5a_check (points: 0)
```
>>> assert get_hash(round(q5a_answer, 5)) == '3401a9c323b4ce1884c91d86e0b0aca6'
```


Tests for q5b (points: 2)
```
>>> assert get_hash(round(q5b_answer, 3)) == 'cee631121c2ec9232f3a2f028ad5c89b'
```


Tests for q6b (points: 1)
```
>>> assert get_hash(round(exceed_budget, 5)) == '579e9eaf15df9f09c9775f131c4d0caf'
```


Tests for q6c (points: 1)
```
>>> assert get_hash(round(q6ci_answer, 5)) == '5274e54635b4e6322c5c33490a61ec94'
```
