```python
# Initialize Otter
import otter
grader = otter.Notebook("ProjPart1.ipynb")
```

# Project - Part 1: Exploring Cook County Housing

## Due Date: Wednesday April 15th 11:59 PM MT on Gradescope

## NO LATE SUBMISSIONS will be accepted - you must plan accordingly.

## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Below are examples of acceptable vs unacceptable use of resources and collaboration when doing the Project assignments in CSCI 3022.


The following would be some **examples of cheating** when working on the Project in CSCI 3022.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given homework problem or one similar is cheating.  However, consulting the class notes, and web pages that explain the material taught in class but do NOT show a solution to the homework problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course notes, the course textbook, and references that contain syntax and/or formulas.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the Project before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a Project problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a Project solution that was your own when asked in office hours by a TA/Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit HW Hours and/or post questions on Piazza!


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Introduction

This project explores what can be learned from an extensive housing dataset that is embedded in a dense social context in Cook County, Illinois.

In project A1 (this assignment), we will guide you through some basic Exploratory Data Analysis (EDA) to understand the structure of the data. Next, you will be adding a few new features to the dataset, while cleaning the data as well in the process. 

In project A2 (the following assignment), you will specify and fit a linear model for the purpose of prediction. Finally, we will analyze the model's error and brainstorm ways to improve its performance.

## Grading
Grading is broken down into autograded answers and free response. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For free response, readers will evaluate how well you answered the question and/or fulfilled the requirements of the question.

Question | Manual | Points
----|----|----
1a | Yes | 4
1b | Yes | 3
2a | No | 3
2b | No | 2
3a | No | 2
3b | Yes | 2
4a | No | 3
4b | No | 2
4c | No | 2
5a | No | 2
5b | No | 4
5c | Yes | 4
6a | No | 2
6b | No | 3
6c | No | 4
6d | No | 2
7a | No | 2
7b | No | 4  
Total |  | 50

```python
import numpy as np

import pandas as pd

%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns

import warnings
warnings.filterwarnings("ignore")

import zipfile
import os

import hashlib

def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()

# Plot settings
plt.rcParams['figure.figsize'] = (12, 9)
plt.rcParams['font.size'] = 12
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# The Data

The dataset consists of over 500,000 records from Cook County, Illinois, the county where Chicago is located. The dataset has 61 features in total; the 62nd is `Sale Price`, which you will predict with linear regression in the next part of this project. An explanation of each variable can be found in the included `codebook.txt` file (you can optionally open this by first clicking the `data` folder, then clicking `codebook.txt` file in the navigation pane). Some of the columns have been filtered out to ensure this assignment doesn't become overly long when dealing with data cleaning and formatting.

The data are split into training and test sets with 204,792 and 68,264 observations, respectively, but we will only be working on the training set for this part of the project.

Let's first extract the data from the `cook_county_data.zip`. Notice we didn't leave the `csv` files directly in the directory because they take up too much space without some prior compression. Just run the cells below: 

```python
with zipfile.ZipFile('data/cook_county_data.zip') as item:
    item.extractall()
```

Let's load the initial data.

```python
initial_data = pd.read_csv("cook_county_train.csv", index_col='Unnamed: 0')
```

As a good sanity check, we should at least verify that the data shape matches the description.

```python
# 204792 observations and 62 features in training data
assert initial_data.shape == (204792, 62)
# Sale Price is provided in the training data
assert 'Sale Price' in initial_data.columns.values
```

The next order of business is getting a feel for the variables in our data.  A more detailed description of each variable is included in `codebook.txt` (in the same directory as this notebook).  **You should take some time to familiarize yourself with the codebook before moving forward.**

Let's take a quick look at all the current columns in our training data.

```python
initial_data.columns.values
```

```python
initial_data['Description'][0]
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: Contextualizing the Data

This dataset was collected by the [Cook County Assessor's Office](https://datacatalog.cookcountyil.gov/Property-Taxation/Archive-Cook-County-Assessor-s-Residential-Sales-D/5pge-nu6u) in order to build a model to predict the monetary value of a home. You can read more about data collection in the CCAO’s [Residential Data Integrity Preliminary Report](https://gitlab.com/ccao-data-science---modeling/ccao_sf_cama_dev/-/blob/master/documentation/Preliminary%20Report%20on%20Data%20Integrity%20June%207,%202019.pdf). 




<!-- BEGIN QUESTION -->

---
## Question 1a

Craft at least two questions about housing in Cook County that could be answered with this data set and provide the type of analytical tool you would use to answer it (e.g. "I would create a ___ plot of ___ and ___" or "I would calculate the ___ [summary statistic] for ___ and ____"). Be sure to reference the columns that you would use and any additional data sets you would need to answer that question.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---
## Question 1b

Suppose now, in addition to the information already contained in the dataset, you also have access to several new columns containing demographic data about the owner, including race/ethnicity, gender, age, annual income, and occupation. Provide one new question about housing in Cook County that can be answered using at least one column of demographic data and at least one column of existing data and provide the type of analytical tool you would use to answer it.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: Exploratory Data Analysis

In Part 2 of this project, you will be building a linear regression model that predicts sales prices using training data, but it's important to first understand how the structure of the data informs such a model. In this section, we will make a series of exploratory visualizations and feature engineering in preparation for that prediction task.

Note that we will perform EDA on the **initial data**.

### Sale Price
We begin by examining the distribution of our target variable `Sale Price`. We have provided the following helper method `plot_distribution` that you can use to visualize the distribution of the `Sale Price` using both the histogram and the box plot at the same time. Run the following 2 cells.

```python
def plot_distribution(data, label):
    fig, axs = plt.subplots(nrows=2)

    sns.distplot(
        data[label], 
        ax=axs[0]
    )
    sns.boxplot(
        x=data[label],
        width=0.3, 
        ax=axs[1],
        showfliers=False,
    )

    # Align axes
    spacer = np.max(data[label]) * 0.05
    xmin = np.min(data[label]) - spacer
    xmax = np.max(data[label]) + spacer
    axs[0].set_xlim((xmin, xmax))
    axs[1].set_xlim((xmin, xmax))

    # Remove some axis text
    axs[0].xaxis.set_visible(False)
    axs[0].yaxis.set_visible(False)
    axs[1].yaxis.set_visible(False)

    # Put the two plots together
    plt.subplots_adjust(hspace=0)
    fig.suptitle("Distribution of " + label)
```

```python
plot_distribution(initial_data, label='Sale Price')
```

We can see that the data is highly skewed to the right currently, probably because of some extremely expensive houses. To get a clearer view of the distribution, run the following cell to examine a more limited range of house prices from [\\$0, \\$1,000,000]. 




```python
no_right_outliers = initial_data[initial_data['Sale Price'] <= 1000000]['Sale Price']
plt.hist(no_right_outliers, bins=200)
plt.xlabel("Sale Price")
plt.ylabel("Count")
plt.title("Distribution of Sale Price")
plt.show()
```

Interestingly, we also see this spike near in the smallest bin of our `Sale Price` distribution, something that wasn't immediately evident from our initial plot. Run the following cell to see which values are the most common.

```python
initial_data['Sale Price'].value_counts()
```

This should be immediately concerning. While we don't know the exact reason why so many houses were reported as \\$1 (most likely a placeholder value of some sort), including these data points in any model would heavily skew the fit. After all, most houses are worth more than \\$1!

<br/>
<hr style="border: 1px solid #fdb515;" />

## Question 2ai

Our findings with the \\$1 entries indicates something interesting may be happening with small valued `Sale Price`s. 

In the cell below, use `value_counts()` to assess the count of entries for `Sale Prices` in the range of [0, 500) (in units dollars). 

Then assign `q2a` to a **list** of the **top 4 most** common `Sale Price`s as integers within the range of [0, 500) in descending order of counts.

```python
...
```

```python
grader.check("q2ai")
```

This exploration has indicated to us that moving forward we should probably pick a cutoff to get rid of smaller, probably unrepresentative `Sale Price`s. 

As a result, **moving forward you will be asked to only look at households whose price is at least \\$500 dollars**. While this cutoff is slightly arbitrary, this decision was made to address these unrealistically low values. 



<br/>
<hr style="border: 1px solid #fdb515;" />

## Question 2aii

On the other hand, there are also the upper-bound outliers that skewed our data in the first place. 

Find the proportion of how many buildings sold for over \\$1,000,000.

```python
...
```

```python
grader.check("q2aii")
```

Since there's such a small percentage of buildings that have sold over \\$1,000,000 it would definitely be reasonable to build a model for the [\\$500, \\$1,000,000] range. 

For this first part of the project though, we're going to try and build a model to handle the entire range from [\\$500, infinity). As a result, we will need to reason with how the outliers shape the distribution of points. 

One way we can make our initial visualization more legible is by log-transforming our data to neutralize the impact of the right outliers.

<br/>
<hr style="border: 1px solid #fdb515;" />

## Question 2b

As previously established, to zoom in on the visualization of most households, we will focus only on a subset of `Sale Price` for this assignment. In addition, it may be a good idea to apply log transformation to `Sale Price`. In the cell below, reassign `training_data` to a new dataframe that is the same as `initial_data` **except with the following changes**:

- `training_data` should contain only households whose price is at least $500.
- `training_data` should contain a new `Log Sale Price` column that contains the log-transformed sale prices.

**You should NOT remove the original column `Sale Price` as it will be helpful for later questions.** If you accidentally remove it, just restart your kernel and run the cells again.

**Note**: This also implies from now on, our target variable in the model will be the log-transformed sale prices from the column `Log Sale Price`. 

*To ensure that any error from this part does not propagate to later questions, there will be no hidden tests for this question.*



```python
...
```

```python
grader.check("q2b")
```

Let's create a new distribution plot using the log-transformed sale prices. As a sanity check, you should see that the distribution for `Log Scale Price` is much more uniform.

```python
plot_distribution(training_data, label='Log Sale Price');
```

<br><br>

<br/>
<hr style="border: 1px solid #fdb515;" />

## Question 3a

Next, we want to explore if there is any correlation between `Log Sale Price` and the total area occupied by the property. The `codebook.txt` file tells us the column `Building Square Feet` should do the trick -- it measures "(from the exterior) the total area, in square feet, occupied by the building".

Let's also apply a log transformation to the `Building Square Feet` column.

In the following cell, create a new column `Log Building Square Feet` in our `training_data` that contains the log-transformed area occupied by each property. 

**You should NOT remove the original `Building Square Feet` column, as it will be used for later questions**. If you accidentally remove it, just restart your kernel and run the cells again.

*To ensure that any errors from this part do not propagate to later questions, there will be no hidden tests for this question.*


```python
training_data['Log Building Square Feet'] = ...
```

```python
grader.check("q3a")
```

<!-- BEGIN QUESTION -->

<!-- BEGIN QUESTION -->

<br/>
<hr style="border: 1px solid #fdb515;" />
         
## Question 3b

In the visualization below, we created a `jointplot` with `Log Building Square Feet` on the x-axis, and `Log Sale Price` on the y-axis. In addition, we fit a simple linear regression line through the bivariate scatter plot in the middle.

Based on the following plot, would `Log Building Square Feet` make a good candidate as one of the features for our model? Why or why not?

**Hint:** To help answer this question, ask yourself: what kind of relationship does a “good” feature share with the target variable we aim to predict?




<img src="images/q2p3_jointplot.png">
    

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/>
<hr style="border: 1px solid #fdb515;" />

## Question 4a

Continuing from the previous part, as you explore the dataset, you might still run into more outliers that prevent you from creating a clear visualization or capturing the trend of the majority of the houses. 

Write a function `remove_outliers` that removes outliers from the dataset based off a threshold value of a variable. For example, `remove_outliers(training_data, 'Building Square Feet', lower=500, upper=8000)` should return a **copy** of `training_data` with only observations that satisfy `Building Square Feet` less than 8000 (exclusive) and `Building Square Feet` greater than 500 (exclusive).

**Note:** The provided tests in this notebook simply check that the `remove_outliers` function you defined does not mutate the input data inplace. They do not check that you have implemented `remove_outliers` correctly so that it works with any data, variable, lower, and upper bound.  That will be tested in hidden tests when you upload to Gradescope.

<!-- grade_id: cell-9186ec2ca053d0aa, points: 0, grade: false, solution: true -->
```python
def remove_outliers(data, variable, lower=-np.inf, upper=np.inf):
    """
    Input:
      data (data frame): the table to be filtered
      variable (string): the column with numerical outliers
      lower (numeric): observations with values lower than this will be removed
      upper (numeric): observations with values higher than this will be removed
    
    Output:
      a data frame with outliers removed
      
    Note: This function should not change mutate the contents of data.
    """  
    ...
```

```python
grader.check("q4a")
```

<br/>
<hr style="border: 1px solid #fdb515;" /> 

## Question 4b 

To give an example on how `remove_outliers` might be useful in the future, plot the distribution of the attribute `Estimate (Land)` below using `plot_distribution` with the provided code below. From `codebook.txt`, `Estimate (Land)` estimates the market value of a property’s land from the prior tax year. 

Indicate if the following statement is correct. Assign your answer to `q4bstatement`. 

“At least 95% of the properties in `training_data` have `Estimate (Land)` values less than $500,000.00”

**Note:** The provided test for this question does not confirm that you have answered correctly; only that you have assigned each variable to `True` or `False`.


```python
# Run this to see the distribution!
plot_distribution(training_data, label='Estimate (Land)');
```

<!-- grade_id: cell-9186ec2ca053d0aa, points: 0, grade: false, solution: true -->
```python
# This should be set to True or False
q4bstatement = ...
```

```python
grader.check("q4b")
```

<br/>
<hr style="border: 1px solid #fdb515;" />

## Question 4c 

Similar to the exploration of `Sale Price` in Question 2, you may have observed that `Estimate (Land)` contains some points that heavily skew its distribution. One way to address these issues is to transform the data to a different scale, as we did with `Sale Price`. Another way is to remove outliers. 

In this subpart, we define an outlier as any point that is 1.5 interquartile ranges above Q3 or below Q1 (exclusive of the boundary). 

 - Assign `IQR` to the interquartile range of `Estimate (Land)` in the initial data. 

 - Then, assign `q4c_training_data` to `training_data` with only rows which are not considered `Estimate (Land)` outliers using `remove_outliers`.

If you need a refresher on how the IQR works, you can read this [Wikipedia page](https://en.wikipedia.org/wiki/Interquartile_range).

**Hint:** The `pandas` function `quantile` [documentation](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.quantile.html) will be useful in calculating the IQR and outliers. 

```python

...



IQR = ...



q4c_training_data = ...
```

```python
grader.check("q4c")
```

Let's create a new distribution plot using the `Estimate (Land)` with outliers removed. As a sanity check for Question 4c, you should see that the distribution for `Estimate (Land)` is much less skewed with `q4c_training_data`!

```python
plot_distribution(q4c_training_data, label='Estimate (Land)');
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 3: Feature Engineering

In this section, we will walk you through a few feature engineering techniques. 

### Bedrooms

Let's start simple by extracting the total number of bedrooms as our first feature for the model. You may notice that the `Bedrooms` column doesn't actually exist in the original `DataFrame`! Instead, it is part of the `Description` column.

<br><br>


    
<hr style="border: 1px solid #fdb515;" />

## Question 5a

Let's take a closer look at the `Description` column first. Compare the description across a few rows together at the same time. For the following list of variables, how many of them can be extracted from the `Description` column? Assign your answer to a list of integers corresponding to the statements that you think are true (ie. `[1, 2, 3]`).

1. The date the property was sold on.
2. The number of stories the property contains.
3. The previous owner of the property.
4. The address of the property.
5. The number of garages the property has.
6. The total number of rooms inside the property.
7. The total number of bedrooms inside the property.
8. The total number of bathrooms inside the property.

```python
# optional cell for scratch work 
```

```python
q5a = ...
```

```python
grader.check("q5a")
```

<br><br>

---

**Regex**


A **regular expression** (“RegEx”) is a sequence of characters that specifies a search pattern. They are written to extract specific information from text. Regular expressions are essentially part of a smaller programming language embedded in python, made available through the re module. As such, they have a stand-alone syntax and methods for various capabilities.

There are a ton of resources to learn and experiment with regular expressions. A few are provided below:

 - [Regex Reference Sheet](https://ds100.org/sp22/resources/assets/hw/regex_reference.pdf)
 - [Official Regex Guide](https://docs.python.org/3/howto/regex.html)
 - [Regex101.com](https://regex101.com/)
  - Be sure to check Python under the category on the left.


Below we work through an example of how to use `regex` to extract the total number of bathrooms from the `Description` column.    

Then you'll apply this knowledge to write your own function to add a column to the dataset with the total number of bedrooms (**as an integer**) for each house.


```python
# Print out Description column and notice where/how bathrooms are saved:
print(training_data["Description"][1])
print("")
print(training_data["Description"][60])
```

Notice the number of bathrooms always shows up in the following statement "and x.y of which are bathrooms." 

Wait, why are these floating point numbers?  Here's how full versus half bathrooms are defined:
 - A full bathroom consists of four main elements — a sink, toilet, shower, and bathtub. 
 - A half-bath is a bathroom that contains two of the four elements, usually a sink and toilet, though in some cases a half-bath could just contain a shower and tub. 


To extract "x.y" as a float from the description column, we can use the following Regular Expression:


`bathrooms_regex = r'and (\d+\.\d) of which are bathrooms'`


 - The `r` preceding the regular expression pattern specifies the regular expression is a raw string. Raw strings do not recognize escape sequences (i.e., the Python newline metacharacter \n). This makes them useful for regular expressions, which often contain literal \ characters. In other words, don’t forget to tag your RegEx with an r.


 - The `(\d+)` is a Regular Expression pattern:
   -  The ( ) specifies the group that you want to capture.
 
   - `\d` is a regex pattern for digit (0-9)

   - `+` is a regex pattern for at least (one or more) 
 
   - `\.` is a decimal point
   
   - We include a second `\d` since we want to grab the number after the decimal point (that is whether there are any 1/2 bathrooms as well)
 
   
 
 
 We can then use [str.extract](https://pandas.pydata.org/docs/reference/api/pandas.Series.str.extract.html) to output this information:
 

```python
# Copy the first 10 rows to use as an example set
example_df = training_data[["Description"]][:10].copy()

# Create a Regex with the format of the bathroom information
bathrooms_regex = r'and (\d+\.\d) of which are bathrooms.'

# Extract the bathroom information from the "Description" column
example_output = example_df["Description"].str.extract(bathrooms_regex).astype(float)

example_output


```


<br><br>

<br/>
<hr style="border: 1px solid #fdb515;" /> 

## Question 5b

Now it's your turn: 

Write a function `add_total_bedrooms(data)` that returns a copy of `data` with an additional column called `Bedrooms` that contains the **total number of bedrooms (as integers**) for each house. Treat missing values as zeros, if necessary. Remember that you can make use of vectorized code here; you shouldn't need any `for` statements. 

**Hint**: You should consider inspecting the `Description` column to figure out if there is any general structure within the text. Once you have noticed a certain pattern, you are set with the power of Regex!

```python
def add_total_bedrooms(data):
    """
    Input:
      data (data frame): a data frame containing at least the Description column.
    Output:
      a Dataframe with a new column "Bedrooms" containing ints.
    
    """
    with_rooms = data.copy()

    ...
    
    return with_rooms

training_data = add_total_bedrooms(training_data)
```

```python
grader.check("q5b")
```

<!-- BEGIN QUESTION -->

<!-- BEGIN QUESTION -->

<br><br>

<br/>
<hr style="border: 1px solid #fdb515;" />
         
## Question 5c

Create a visualization that clearly and succinctly shows if there exists an association between  `Bedrooms` and `Log Sale Price`. A good visualization should satisfy the following requirements:
- It should avoid overplotting.
- It should have clearly labeled axes and a succinct title.
- It should convey the strength of the correlation between `Sale Price` and the number of rooms: in other words, you should be able to look at the plot and describe the general relationship between `Log Sale Price` and `Bedrooms`

**Hint 1**: A direct scatter plot of the `Sale Price` against the number of bedrooms for all of the households in our training data might risk overplotting. Try it out and see why causes overplotting. Then choose a different visualization method to avoid overplotting.

**Hint 2**: Take a closer look at the variable `Bedrooms`. What variable type is `Bedrooms` and what are the possible values `Bedrooms` can take? What visualizations might be good at showing this?

```python
...
```

<!-- END QUESTION -->

<br><br>

<br/>
<hr style="border: 1px solid #fdb515;" />

## Question 6


Now, let's take a look at the relationship between neighborhood and sale prices of the houses in our dataset.
Notice that currently we don't have the actual names for the neighborhoods. Instead we will use a similar column `Neighborhood Code` (which is a numerical encoding of the actual neighborhoods by the Assessment office).


## Question 6a

Before creating any visualization, let's quickly inspect how many different neighborhoods we are dealing with.

Assign the variable `num_neighborhoods` with the total number of unique neighborhoods in `training_data`. 


```python
num_neighborhoods = ...
num_neighborhoods
```

```python
grader.check("q6a")
```

<br><br>

---
## Question 6b

If we try directly plotting the distribution of `Log Sale Price` for all of the households in each neighborhood using the `plot_categorical` function from the next cell, we would get the following visualization.



<img src="images/q5p2_catplot.png">

```python
# Feel free to use the cell below this and run plot_categorical(training_data) if you want to see what this function outputs.

def plot_categorical(neighborhoods):
    fig, axs = plt.subplots(nrows=2)

    sns.boxplot(
        x='Neighborhood Code',
        y='Log Sale Price',
        data=neighborhoods,
        ax=axs[0],
    )

    sns.countplot(
        x='Neighborhood Code',
        data=neighborhoods,
        ax=axs[1],
    )

    # Draw median price
    axs[0].axhline(
        y=training_data['Log Sale Price'].median(), 
        color='red',
        linestyle='dotted'
    )

    # Label the bars with counts
    for patch in axs[1].patches:
        x = patch.get_bbox().get_points()[:, 0]
        y = patch.get_bbox().get_points()[1, 1]
        axs[1].annotate(f'{int(y)}', (x.mean(), y), ha='center', va='bottom')

    # Format x-axes
    axs[1].set_xticklabels(axs[1].xaxis.get_majorticklabels(), rotation=90)
    axs[0].xaxis.set_visible(False)

    # Narrow the gap between the plots
    plt.subplots_adjust(hspace=0.01)
```

```python

```

Oh no, looks like we have run into the problem of overplotting again! 

You might have noticed that the graph is overplotted because **there are actually quite a few neighborhoods in our dataset**! For the clarity of our visualization, we will have to zoom in again on a few of them. The reason for this is our visualization will become quite cluttered with a super dense x-axis.

Assign the variable `top_15_neighborhood_codes` to a **list** of the 15 neighborhood codes that have the greatest number of property sales within them.


Assign the variable `in_top_15_neighborhoods` to a copy of `training_data` that has been filtered to only contain rows corresponding to properties that are in one of the top 15 most populous neighborhoods. 


```python

top_15_neighborhood_codes = ...

in_top_15_neighborhoods = ...


```

```python
grader.check("q6b")
```

Let's create another of the distribution of sale price within in each neighborhood again, but this time with a narrower focus!

```python
plot_categorical(neighborhoods=in_top_15_neighborhoods)
```

<br><br>

---
## Question 6c

From the plot above, we can see that there is much less data available for some neighborhoods. For example, Neighborhood 81 has only around 33% of the number of datapoints as Neighborhood 30.

One way we can deal with the lack of data from some neighborhoods is to create a new feature that bins neighborhoods together. We’ll categorize our neighborhoods in a crude way. In Question 6c, we’ll compute how “expensive” each neighborhood is by aggregating the `Log Sale Price`s for all properties in a particular neighborhood using a `metric`, such as the median. We’ll use this `metric` to find the top `n` most expensive neighborhoods. Then, in Question 6d, we’ll label these “expensive neighborhoods” and leave all other neighborhoods unmarked.

Fill in `find_expensive_neighborhoods` to return a **list** of the neighborhood codes of the **top `n`** most expensive neighborhoods as measured by our choice of aggregating function, `metric`.

For example, calling `find_expensive_neighborhoods(training_data, n=3, metric=np.median)` should return the 3 neighborhood codes with the highest **median `Log Sale Price`** computed across all properties in those neighborhood codes. 

Note that there the tests below just check that you have the correct format, NOT whether or not you have the correct answer.  That will be checked via a hidden test in Gradescope.


```python
def find_expensive_neighborhoods(data, n=3, metric=np.median):
    """
    Input:
      data (data frame): should contain at least a string-valued 'Neighborhood Code'
        and a numeric 'Sale Price' column
      n (int): the number of top values desired
      metric (function): function used for aggregating the data in each neighborhood.
        for example, np.median for median prices
    
    Output:
      a list of the the neighborhood codes of the top n highest-priced neighborhoods as measured by the metric function
    """
    neighborhoods = ...
    
    # This makes sure the final list contains the generic int type used in Python3, not specific ones used in numpy.
    return [int(code) for code in neighborhoods]

expensive_neighborhoods = find_expensive_neighborhoods(training_data, 3, np.median)
expensive_neighborhoods
```

```python
grader.check("q6c")
```

<br><br>

---
## Question 6d

We now have a list of neighborhoods we've deemed as higher-priced than others.  Let's use that information to write an additional function `add_expensive_neighborhood` that takes in a `DataFrame` of housing data (`data`) and a list of neighborhood codes considered to be expensive (`expensive_neighborhoods`). You can think of `expensive_neighborhoods` as being the output of the function `find_expensive_neighborhoods` from Question 6c. 

Using these inputs, `add_expensive_neighborhood` should add a column to `data` named `in_expensive_neighborhood` that takes on the **integer** value of 1 if a property is part of a neighborhood in `expensive_neighborhoods` and the integer value of 0 if it is not. This type of variable is known as an **indicator variable**.

**Hint:** [`pd.Series.astype`](https://pandas.pydata.org/pandas-docs/version/0.23.4/generated/pandas.Series.astype.html) may be useful for converting `True`/`False` values to integers.


```python
def add_in_expensive_neighborhood(data, expensive_neighborhoods):
    """
    Input:
      data (DataFrame): a DataFrame containing a 'Neighborhood Code' column with values
        found in the codebook
      expensive_neighborhoods (list of ints): ints should be the neighborhood codes of
        neighborhoods pre-identified as expensive
    Output:
      DataFrame identical to the input with the addition of a binary
      in_expensive_neighborhood column
    """
    data['in_expensive_neighborhood'] = ...
    return data

expensive_neighborhoods = find_expensive_neighborhoods(training_data, 3, np.median)
training_data = add_in_expensive_neighborhood(training_data, expensive_neighborhoods)


```

```python
grader.check("q6d")
```

<br><br>

In the following question, we will take a closer look at the `Roof Material` feature of the dataset and examine how we can incorporate categorical features into the linear models we will be developing in Part 2 of the project.

---
## Question 7a

If we look at `codebook.txt` carefully, we can see that the Assessor's Office uses the following mapping for the numerical values in the `Roof Material` column.
```
Roof Material (Nominal): 

       1    Shingle/Asphalt
       2    Tar&Gravel
       3    Slate
       4    Shake
       5    Tile
       6    Other
```

**NOTE** that Tar&Gravel has no spaces. That is, 'Tar&Gravel' $\neq$ 'Tar & Gravel'

Write a function `substitute_roof_material` that replaces each numerical value in `Roof Material` with a `string` of the corresponding roof material. Your function should return a new `DataFrame`, not modify the existing `DataFrame`.  If you modify the existing `DataFrame` by accident, you can load `training_data` again in `q2b`.

**Hint**: the `DataFrame.replace` ([documentation](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.replace.html)) method may be useful here.





```python
def substitute_roof_material(data):
    """
    Input:
      data (DataFrame): a DataFrame containing a 'Roof Material' column.  Its values
                         should be limited to those found in the codebook
    Output:
      new DataFrame identical to the input except with a refactored 'Roof Material' column
    """
    ...
    return data
    
training_data_mapped = substitute_roof_material(training_data) 


training_data_mapped["Roof Material"].unique()
```

```python
grader.check("q7a")
```

---
## Question 7b

#### One Hot Encoding 

Unfortunately, simply replacing the integers with the appropriate strings isn’t sufficient for using `Roof Material` in our model.  Since `Roof Material` is a categorical variable, we will have to one-hot-encode the data. For more information on why we want to use one-hot-encoding, refer to this [link](https://machinelearningmastery.com/why-one-hot-encode-data-in-machine-learning/).

Complete the following function `ohe_roof_material` that returns a copy of the original dataframe with the new column one-hot-encoded on the roof material of the household. These new columns should have the form `Roof Material_MATERIAL`.  In this problem we only wish to construct the one-hot-encoding columns **without removing any columns**. Your function should return **a new DataFrame, not modify the existing DataFrame.**  

**Note**: You should **avoid using `pd.get_dummies`** in your solution as it will remove your original column and is therefore not as reusable as your constructed data preprocessing pipeline. 
Instead should use `scikit-learn`’s [OneHotEncoder]((https://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.OneHotEncoder.html))
to perform the one-hot-encoding. `OneHotEncoder` will automatically generate column names of the form `Roof Material_MATERIAL`



*Hint*: To get you started with this subpart, the next few boxes of code below initializes a `OneHotEncoding` pre-processing "model" from Scikit-learn and fits it on a simple dataset containing random names. Please play with this code before jumping into the roof material data if you are unsure how to approach the question using `OneHotEncoder`.


```python
# Sample code to demonstrate One Hot Encoding

d = {'Names': ['Anirudhan', 'Dominic', 'Rahul', 'Rahul', 'Anirudhan', 'Yike', 'Vasanth']}
demo_df = pd.DataFrame(data=d)

demo_df

```

```python
# Initialize the one hot encoder for sample demo
from sklearn.preprocessing import OneHotEncoder

ohe = OneHotEncoder()
```

To "learn" the categories we fit the OneHotEncoder to the categorical data:

```python
ohe.fit(demo_df[["Names"]])
```

We can get the "names" of the new one-hot-encoding columns which reveal both the source columns and the categories within each column: 

```python
ohe.get_feature_names_out()
```

We can also construct the OneHotEncoding of the data:

```python
ohe.transform(demo_df[["Names"]])
```

Notice that the One-Hot-Encoding produces a sparse output matrix.  This is because most the entries are 0.    To view this matrix, we need to convert it to a numpy matrix using the following syntax:



```python
encoded_names = ohe.transform(demo_df[["Names"]]).toarray()

encoded_names
```

We can then convert this to a DataFrame (to merge with our original demo_df DataFrame) as follows:

```python

encoded_demo_df = pd.DataFrame(encoded_names, columns=ohe.get_feature_names_out(), index = demo_df.index)

encoded_demo_df

```

```python
#Now join back with original df:

ohe_demo = demo_df.join(encoded_demo_df)

ohe_demo
```

Now it's your turn! 
Complete the following function `ohe_roof_material` that returns a copy of the original dataframe with the new column one-hot-encoded on the roof material of the household. These new columns should have the form `Roof Material_MATERIAL`.  In this problem we only wish to construct the one-hot-encoding columns **without removing any columns**. Your function should return **a new DataFrame, not modify the existing DataFrame.**  



```python
from sklearn.preprocessing import OneHotEncoder

def ohe_roof_material(data):
    """
    One-hot-encodes roof material.  New columns are of the form "Roof Material_MATERIAL"
    """
    ...

training_data_ohe = ohe_roof_material(training_data_mapped)

#This line of code will display only the one-hot-encoded columns in training_data_ohe that 
# have names that begin with “Roof Material_" 

training_data_ohe.filter(regex='^Roof Material').head(5)
```

```python
grader.check("q7b")
```

## Takeaways

In this project, you have...
- Peformed basic Exploratory Data Analysis (EDA)
- Added new features to the data


## Congratulations! You have finished the Project Part 1!

### What's next? 

In Project A2, you will focus on building a linear regression model to predict house prices. You are well-prepared to build such a model: you have considered what is in this dataset, what it can be used for, and engineered some features that should be useful for prediction. 

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any websites you referenced

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 

* **Project Part 1 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the Project Part 1 Autograded assignment in Gradescope.

* **Project Part 1 Manually Graded**: Submit your ProjPart1.PDF to the HW10 Manually Graded assignment in Gradescope.  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.  IF NOT, YOU WILL LOSE POINTS**   Also, **check** that all of your plots **and** all lines of your code are showing up in your PDF before submitting.  If not, you will not receive credit for your plots/code.  


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

```python
import simple_latex_checker as slc

nb = slc.Nb_checker()
nb.run_check("ProjPart1.ipynb")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='ProjPart1.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q2ai (points: 2)
```
>>> assert isinstance(q2ai, list) and len(q2ai) >= 3 and isinstance(q2ai[0], int)
```
```
>>> assert get_hash(q2ai[0]) == 'c4ca4238a0b923820dcc509a6f75849b'
```
```
>>> assert get_hash(q2ai[1]) == 'f899139df5e1059396431415e770c6dd'
```
```
>>> assert get_hash(q2ai[2]) == 'd3d9446802a44259755d38e6d163e820'
```
```
>>> assert get_hash(q2ai[3]) == '38b3eff8baf56627478ec76a704e9b52'
```


Tests for q2aii (points: 1)
```
>>> assert get_hash(round(q2aii, 5)) == '8783110f5dbe08cb143cdb7b3e344db3'
```


Tests for q2b (points: 2)
```
>>> assert ('Sale Price' in training_data.columns) == True and ('Log Sale Price' in training_data.columns) == True
```
```
>>> assert get_hash(round(training_data['Log Sale Price'].sum(), 3)) == '7781705cc447cc5af213156282d855fd'
```
```
>>> assert training_data.shape[0] == 168931
```


Tests for q3a (points: 2)
```
>>> assert get_hash(round(training_data['Log Building Square Feet'].mean(), 3)) == 'a9af50facbc916992c1293e018cb4ec3'
```


Tests for q4a (points: [0, 0, 1, 3])
```
>>> assert remove_outliers(training_data, 'Building Square Feet', upper=2000).shape != training_data.shape
```
```
>>> assert np.isclose(sum(training_data['Building Square Feet']), 274753144.0, atol=0.1) == True
```
```
>>> assert np.isclose(sum(remove_outliers(training_data, 'Sale Price', lower=121000, upper=1000000)['Building Square Feet']), 203454261.0, atol=0.1)
```
```
>>> assert remove_outliers(training_data, 'Estimate (Building)', lower=0, upper=8187420).shape[0] == 167693
```


Tests for q4b (points: [0])
```
>>> assert set([q4bstatement]).issubset({False, True})
```


Tests for q4c (points: 2)
```
>>> assert get_hash(int(IQR)) == 'd2d1836b7375e85336296d909745a1ad'
```
```
>>> assert q4c_training_data.shape[0] >= 120000 and q4c_training_data.shape[0] <= 160000
```
```
>>> assert get_hash(int(q4c_training_data['Estimate (Land)'].max())) == '9117ad1cb16279cfce2a239e12e4c66c'
```
```
>>> assert get_hash(int(q4c_training_data['Estimate (Land)'].min())) == 'cfcd208495d565ef66e7dff9f98764da'
```


Tests for q5a (points: [0, 2])
```
>>> assert type(q5a) == list
```
```
>>> assert set([1, 2, 4, 6, 7, 8]) == set(q5a)
```


Tests for q5b (points: 4)
```
>>> assert not training_data['Bedrooms'].isnull().any()
```
```
>>> assert get_hash(training_data['Bedrooms'].sum()) == '4ee51f0a44558acba5d037b211ae878a'
```
```
>>> assert get_hash(training_data.loc[10000, 'Bedrooms']) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```
```
>>> assert get_hash(sum(training_data.loc[:, 'Bedrooms'] < 5)) == '44ff361eb96953006ac4188e9c3ccc2a'
```


Tests for q6a (points: 2)
```
>>> assert get_hash(num_neighborhoods) == 'bd686fd640be98efaae0091fa301e613'
```


Tests for q6b (points: [0, 0, 0, 2, 1])
```
>>> assert get_hash(len(in_top_15_neighborhoods['Neighborhood Code'].unique())) == '9bf31c7ff062936a96d3c8bd1f8f2ff3'
```
```
>>> assert isinstance(top_15_neighborhood_codes, list)
```
```
>>> assert get_hash(in_top_15_neighborhoods['Neighborhood Code'].iloc[0]) == 'da4fb5c6e93e74d3df8527599fa62642'
```
```
>>> assert get_hash(sum(top_15_neighborhood_codes)) == 'db576a7d2453575f29eab4bac787b919'
```
```
>>> assert get_hash(in_top_15_neighborhoods['Neighborhood Code'].sum()) == 'a691bdb9bae6abe8b64fbe61b18b8726'
```


Tests for q6c (points: [0, 0])
```
>>> assert get_hash(len(find_expensive_neighborhoods(training_data, 5, np.median))) == 'e4da3b7fbbce2345d7772b0674a318d5'
```
```
>>> assert isinstance(expensive_neighborhoods, list)
```


Tests for q6d (points: [0, 2])
```
>>> assert sum(training_data.loc[:, 'in_expensive_neighborhood']) == 1290 and sum(training_data.loc[:, 'in_expensive_neighborhood'].isnull()) == 0
```
```
>>> assert get_hash(sum(add_in_expensive_neighborhood(training_data, [44, 94, 93]).loc[:, 'in_expensive_neighborhood'])) == '70222949cc0db89ab32c9969754d4758'
```


Tests for q7a (points: [0, 2])
```
>>> assert set(training_data_mapped['Roof Material'].unique()) == set(['Shingle/Asphalt', 'Tar&Gravel', 'Other', 'Tile', 'Shake', 'Slate'])
```
```
>>> assert training_data_mapped['Roof Material'].value_counts().values.tolist() == [160760, 3774, 1620, 1201, 841, 735]
```


Tests for q7b (points: [1, 1, 2])
```
>>> assert training_data_ohe.shape == (168931, 72)
```
```
>>> expected_ohe_cols = ['Roof Material_Other', 'Roof Material_Shake', 'Roof Material_Shingle/Asphalt', 'Roof Material_Slate', 'Roof Material_Tar&Gravel', 'Roof Material_Tile']
>>> ohe_values = training_data_ohe[expected_ohe_cols].to_numpy()
>>> assert ohe_values.sum() == training_data_ohe.shape[0]
```
```
>>> expected_ohe_cols = set(['Roof Material_Other', 'Roof Material_Shake', 'Roof Material_Shingle/Asphalt', 'Roof Material_Slate', 'Roof Material_Tar&Gravel', 'Roof Material_Tile'])
>>> assert expected_ohe_cols.issubset(set(training_data_ohe.columns)) == True
```
