```python
# Initialize Otter
import otter
grader = otter.Notebook("hw08.ipynb")
```

# Homework 8:  Central Limit Theorem and Confidence Intervals



## Due Date: Wednesday March 11th by 11:59 PM MT on Gradescope


## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Below are examples of acceptable vs unacceptable use of resources and collaboration when doing HW assignments in CSCI 3022.


The following would be some **examples of cheating** when working on HW assignments in CSCI 3022.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given homework problem or one similar is cheating.  However, consulting the class notes, and web pages that explain the material taught in class but do NOT show a solution to the homework problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course notes, the course textbook, and references that contain syntax and/or formulas.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the HW before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by a TA/Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit HW Hours and/or post questions on Piazza!


#### Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1 | 3 | autograded |
Question 2 | 8 | manual |
Question 3a,c,d | 8 | autograded |
Question 3b,e | 6 | manual |
Question 4 | 8 | autograded |
Question 5 | 9 | autograded |
Question 6a,c,d | 4 | autograded |
Question 6b | 2 | manual & autograded |
Question 6e | 2 | manual |
|Total | 50



**Assigned Readings On These Topics from Class Schedule**: 

* [The Central Limit Theorem](https://inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html?highlight=central+limit+theorem)
* [Variability of the Sample Mean](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html)

* [Confidence Intervals](https://inferentialthinking.com/chapters/13/3/Confidence_Intervals.html)

* [Using Confidence Intervals](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html)

* [Choosing a Sample Size](https://inferentialthinking.com/chapters/14/6/Choosing_a_Sample_Size.html)

* [Bootstrapping](https://inferentialthinking.com/chapters/13/2/Bootstrap.html)



**Assignment Learning Objectives**

* Explain the difference between parameters and statistics 

* Explain the difference between a population distribution, a sample distribution and the sampling distribution of a statistic

* Simulate the empirical sampling distribution of a statistic

* Apply and interpret the Central Limit Theorem for sums and means (what is it used for?  What conditions must hold to use it?)

* Define what a confidence interval is and how to interpret it.  

* Calculate confidence intervals using the bootstrapping percentile method.

* Explain key assumptions and limitations of the bootstrapping percentile method.

* Use the CLT to create and interpret Confidence Intervals for the population mean.

* Use CLT to create and interpret confidence intervals for population proportions.

* Determine sample sizes based on desired confidence interval widths.



## <a id='top'></a>
---
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## IMPORTANT STEP BEFORE STARTING THIS ASSIGNMENT. 

To make sure all of your plots render correctly in the PDF, please do the following steps now: 

1.) In the toolbar, go to Settings 

2.) Disable Autosave Documents 

```python
import pandas as pd
# These lines import the Numpy and Datascience modules.
import numpy as np
from matplotlib.ticker import PercentFormatter
%matplotlib inline
# These lines do some fancy plotting magic.
import matplotlib

import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')

import hashlib

from matplotlib import patches
#from ipywidgets import interact, interactive, fixed
#import ipywidgets as widgets


def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

***


<a id='p2'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴 Question 1 

A hardware store receives a shipment of 10,000 bolts that are supposed to be 12 cm long. 
The mean of this shipment of 10,000 bolts is indeed 12 cm, and the standard deviation is 0.2 cm.   

For the following questions, determine if you have enough information to answer.  If you do, then use built-in python functions to calculate the probability. If not, enter -1 for your answer.

```python
# Run this cell to import stats
from scipy import stats as stats
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 1a (1 pt)
    
What is the probability that a randomly chosen bolt is less than 10 cm long?

```python
q1a_answer = ...
```

```python
grader.check("q1a")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 1b (2 pts)

For quality control, the hardware store chooses 100 bolts at random to measure.  They will declare the shipment defective and return it to the manufacturer if the average length of 100 bolts is less than 11.97 cm or greater than 12.04 cm.    Find the probability that the shipment is found satisfactory (i.e. not defective).

```python
q1b_answer = ...
```

```python
grader.check("q1b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 2  ##

The file `data/income_data.csv` contains Age and Income information from a fictitious city in the Midwest with 5000 residents.   Consider this data to be our population data.

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 2a (2 pts) ## 

Load the data into a pandas DataFrame called `dfIncome`, calculate the mean of the Income data and then  make a **density histogram** of the Distribution of the Income data **with 15 bins**.   Include a title for your plot and label the x-axis (we have provided a label for the y-axis).
Note we have included code to mark where the population mean lies on the histogram.

```python
dfIncome = ...

print(np.median(dfIncome["Income"]))

mean_income = ...
print(mean_income)

print("Population income mean is", mean_income)

...
plt.xticks(np.arange(0, 210000, step=50000))
# write code to plot density histogram above this line

plt.ylabel("Probability Per Dollar") # Since this is a density histogram, the y-units are probability per unit x
#Add a triangle marker to indicate where the population mean is
plt.scatter(mean_income, -.0000001, marker='^', color='red', s=300)

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 2bi (2 pts) ## 

Write a function to collect a random sample of size `sample_size` with replacement from `dfIncome`  and plot the **density histogram** of the empirical distribution of the income for your sample.  Use 15 bins for your histogram and set the x-axis range to be from (0,210000).   (Hint: use the dataframe method `.sample()`). Include a title and label for both axes.          
Then run the cells provided below to output 3 separate distributions for sample sizes of $10$, $100$ and $1000$. 


```python
def income_sample(df, sample_size):
    
...
    plt.xticks(np.arange(0, 210000, step=50000))
    
# rest of code above this line
```

```python
income_sample(dfIncome,10)
```

```python
income_sample(dfIncome,100)
```

```python
income_sample(dfIncome,1000)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 2bii (1 pt) ## 

What happens to the shape of the empirical sample distributions of income as you increase the sample size, and how does it relate to the population distribution? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 2c (2 pts) ## 

If we want to estimate the **mean** income of the population we can draw a sample from the population and compute the sample mean.  As we learned in class, since samples can vary, the sample mean can vary and thus it is a random variable and has its own distribution.


i).  Complete the function `income_sample_mean` below to randomly sample `sample_size` rows from `dfIncome` with replacement and return the **sample mean of income** for that sample. 


ii).  Complete the function `income_sample_mean_dist` below to simulate `num_simulations` of randomly sampling `sample_size` rows from `dfIncome` with replacement and calculate the sample mean of income for each sample.
Store the sample means in an np.array called `means`.  The function should output a **density** histogram of the empirical sample mean income distribution.    On the histogram, include two **triangle** markers on the histogram like in part 1a: A red one for the population mean (that you calculated in part 1A) and a yellow one for the mean of the `num_simulations` sample mean estimates.
Use 25 bins, do not set a range, and keep the x-axis limits fixed across all plots for comparison.
Include a title and labels for the x and y axes.

Then run the cells provided below to output 3 separate distributions for `num_simulations`=100,000 and `sample_size` = $10$, $100$ and $1000$



```python
def income_sample_mean(df, sample_size):
    ...
```

```python
def income_sample_mean_dist(df, sample_size, num_simulations):
    means = ...
      # 'means' stores "num_simulations" means from samples of size "sample_size"
    
    ...
    plt.xlim([25000,95000])

    
# Your code for part (ii) above
```

```python
income_sample_mean_dist(dfIncome, 10, 100000)
```

```python
income_sample_mean_dist(dfIncome, 100, 100000)
```

```python
income_sample_mean_dist(dfIncome, 1000, 100000)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 2d (1 pt) ## 

Describe the shapes of the empirical sample mean distributions. 
What happens to the mean and standard deviations of these distributions as you increase the sample size?
What is the name of the theorem that explains what you are observing?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 3 

We rarely have data from an entire population, like we had available in Question 2.  Instead we usually only have access to a sample.  

Let `sample` be the following random sample of size 500 with replacement from the `dfIncome["Income"]` data from Question 2.
To ensure we are all working with the same sample, we can set the `random_state` variable to a specific integer (https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.sample.html)



```python
# Just run this cell, so we are all working with the same sample data.
sample = dfIncome["Income"].sample(500, replace=True, random_state = 100)

```

Plot a density histogram of the sample data:

```python
# Plot Density Histogram of sample data:


  
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 3a (3 points)

i). Calculate the sample mean and save it to the variable `sample_mean`

ii). Use the Central Limit Theorem to use this sample data to find an approximate 95% confidence interval for the population mean income. 

Save the lower value of your CI as `CI_lower` and the upper value as `CI_upper`


```python

sample_mean = ...

...


CI_lower = ...

CI_upper = ...

print("95% CI for the population average income is [", round(CI_lower,2), ",", round(CI_upper,2), "]")

```

```python
grader.check("q3a")
```

<!-- BEGIN QUESTION -->


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 3b (4 points)

i). Describe in words how we can correctly interpret this confidence interval.  
What does this interval tell us about the population? 
What does the $95\%$ confidence level mean in this context?

ii).  In this contrived example, we can actually calculate the population average income.  What is the population average income?  Is it within this particular confidence interval or not?  


Write your answer to **part (i)** in this cell. 

```python


# Answers to part (ii)
pop_avg_income = ...

# If the pop_avg_income is actually in the confidence interval you found in part (a), answer with the string "Yes", otherwise answer "No"

pop_avg_in_CI = ...

```

```python
grader.check("q3b")
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 3c (4 points)


This question involves creating confidence intervals using bootstrapping.
If you'd like a refresher of the concepts needed to answer this question, we recommend watching the following:

 -  **[Video: Intro to Confidence Intervals by Professor Oscamou](https://www.youtube.com/watch?v=pCr_Fggyp4w)**
 -  **[Video: Bootstrapping Confidence Intervals by Professor Oscamou](https://www.youtube.com/watch?v=x2NrjZM4fgQ)**
----------------------------------
                                                                       

Suppose instead of a confidence interval for the population average income, we are interested in a confidence interval for the population **median** income.  The median may be a more useful parameter since the Income distribution is skewed. 

In this situation, the Central Limit Theorem does not apply.  Instead, we can use bootstrapping to compute the confidence interval.  
    
i). Complete the function `one_resampled_med` to take an input array of data, conduct one bootstrapped resample of the data and output the sample median of the bootstrapped resample.  **To pass the autograder**, you must use `np.random.choice` to create your resample.

ii).  Use list comprehension to run the function `one_resampled_med` 100,000 times and store the results in a numpy array called `resampled_med`. 

iii). Create a density histogram of the entries in the `resampled_med` array.  Set the bin number equal to 20 when plotting your histogram. Label both axes and include a title on your plot.


```python
def one_resampled_med(data):
    ...

one_resampled_med(sample)
```

```python
resampled_med = ...

# answer to part ii above
```

```python
...
# iii). your code for histogram above this line
```

```python
grader.check("q3c")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 3d (1 point)
Using the array `resampled_med`, find the values at the two edges of the middle 95% of the bootstrapped mean estimates. Compute the lower and upper ends of the interval, named `CI_med_lower` and `CI_med_upper` respectively. 


```python
CI_med_lower = ...
CI_med_upper= ...
print(f"Bootstrapped 95% confidence interval for the population median income: [{CI_med_lower:.2f}, {CI_med_upper:.2f}]")
```

```python
grader.check("q3d")
```


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 3e (1 pt)

Again, in this contrived example we can calculate the actual population median income.  What is it?  Is it within the 95% CI interval you obtained via bootstrapping? 
 


```python


pop_median_income = ...

# If the pop_med_income is actually in the confidence interval you found in part (f), answer with the string "Yes", otherwise answer "No"

pop_median_in_CI = ...

```

```python
grader.check("q3e")
```

<br>

---

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 4 (8 pts)


#### Video Assignment

Complete this [Video Assignment: Confidence Intervals for Proportions (click this link)](https://canvas.colorado.edu/courses/133218/assignments/2588776) available on Canvas (to receive points you must watch the entire video and answer the embedded questions).  See the accompanying slides in Lesson 18.


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

[Back to top](#top)

<a id='p4'></a>

## 🔴  Question 5. Sample Size and Confidence Level


A data science class at the large University wants to estimate the **proportion of students at the school who use Instagram.** To do this, they need to take a random sample of students. You can assume that their method of sampling is equivalent to drawing at random with replacement from students at the school.


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 5a (3 points)
Assign `smallest` to the **integer** (not float) that gives the smallest number of students they should sample to ensure that a **95%** confidence interval for the population proportion that has a width of no more than $8\%$ from left end to right end. 
                                                                                                                                                           
Note that the textbook and video example round 1.96 to 2 when doing examples, but for this question, do not round.                                                                                                                                                            

*Hint:* How can our data be represented to show if a student in the sample is an Instagram user? Given this, what assumptions can we make for the SD of the population? [Section 14.6](https://inferentialthinking.com/chapters/14/6/Choosing_a_Sample_Size.html#) might be helpful!
                                                                                            
                                                                                                                                                                                
*Note:* Notice that the smallest sample size has to be an integer, so if you get a non-integer for an answer, you need to round it up to the nearest integer.  The `np.ceil` function may be helpful.


```python
smallest = ...

```

```python
grader.check("q5a")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 5b (2 points)
    
Suppose the data science class decides to construct a 90% confidence interval instead of a 95% confidence interval, but they still require that the width of the interval is no more than 8% from left end to right end. Will they need the same sample size as in Question 5a? 

Assign `sample_90` to the number corresponding to the correct answer.

1. Yes, they must use the same sample size.
2. No, a smaller sample size will work.
3. No, they will need a bigger sample.


```python
sample_90 = ...
```

```python
grader.check("q5b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 5c (2 points)
The professor tells the class that a 90% confidence interval for the parameter is constructed exactly like a 95% confidence interval, except that you have to go only 1.64 SDs on either side of the estimate (±1.64) instead of 1.96 SDs on either side (±1.96). Assign `smallest_num` to the smallest number of students they should sample to ensure that a **90%** confidence interval for the parameter has a width of no more than 8% from left end to right end. 

*Note:* Just like in Question 5a, your answer should an **integer**. 


```python
smallest_num = ...

```

```python
grader.check("q5c")
```

For this next exercise, please consult [Section 14.3.4](https://inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html#the-standard-normal-cdf) of the textbook for similar examples.

The students are curious about how the professor came up with the value 1.64 in Question 5c. 

To start, run the code below to visualize a plot of the standard normal distribution (with mean 0 and standard deviation 1).  

The `plt.fill_between` method shades in the region from -1.64 to +1.64 SDs. 


```python
# Just run this cell

x = np.linspace(-4, 4, 1000)
plt.plot(x, stats.norm.pdf(x,0))
plt.xticks([-3,-1.64,0,1.64,3])
plt.title("Normal Distribution: mean = 0, std = 1")
plt.fill_between(x, stats.norm.pdf(x,0), where=[(x > -1.64) and (x < 1.64) for x in x], alpha=0.4)

```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 5d (1 point)
To verify that 1.64 is in fact the correct cutoff for a 90% confidence interval, use built-in python functions to calculate the shaded area in the plot above.  Assign your answer to the variable `shaded_area`

*Hint:* Recall the acronym `cdf` stands for `cumulative distribution function`. It measures the proportion to the left of a specified point under a probability histogram.

```python
shaded_area = ...
round(shaded_area, 2)

```

```python
grader.check("q5d")
```

The acronym `ppf` stands for `probability point function`.  It's the inverse of the `cdf`.  

Specifically, `ppf(y)` returns the exact point where the probability of everything to the left is equal to y. This can be thought of as the percentile function since the ppf tells us the value of a given percentile of the data.

To find the lower SD cut-off for a 90% confidence interval, notice that we want the value on the x-axis of the standard normal distribution such that the area to left is equal to $\frac{1-.90}{2}$ = 0.05

We can find this by running the following code:


```python
stats.norm.ppf(.05)
```

By symmetry the upper SD cutoff is the opposite of this value:

```python
-stats.norm.ppf(.05)
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 5e (2 points)

a).  To construct a 99% confidence interval for a parameter, how many Standard Errors do you need to add and subtract to the sample estimate?   Give your answer as a positive value, rounded to the nearest hundredth.  Assign `SE_99` to your answer.  


b). Using your rounded answer for `SE_99` from part a, assign `smallest_num_99` to the **integer** (not float) that gives the smallest number of students you should sample to ensure that a **99%** confidence interval for the parameter has a width of no more than $8\%$ from left end to right end.  



```python
SE_99 = ...

```

```python
smallest_num_99 = ...

```

```python
grader.check("q5e")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

[Back to top](#top)

<a id='p2'></a>

## 🔴  Question 6. Thai Restaurants in Boulder



Oswaldo and Varun are trying to see what the best Thai restaurant in Boulder is. They survey 1,500 CU Boulder students selected uniformly at random and ask each student what Thai restaurant is the best. (*Note: This data is fabricated for the purposes of this homework.*) The choices of Thai restaurants are [Aloy Thai](https://www.aloythai.com/), [Budda Cafe](https://www.buddhathaicuisine.com/menu), [Folsom Thai](https://www.folsomthaiatboulder.com/), and [Khow Thai](https://khow-thai.com/). After compiling the results, Oswaldo and Varun release the following percentages from their sample:

|Thai Restaurant  | Percentage|
|:------------:|:------------:|
|Aloy Thai| 8% |
|Khow Thai | 53% |
|Folsom Thai | 25% |
|Budda Cafe | 14% |

These percentages represent a uniform random sample of the population of CU Boulder students. We will attempt to estimate the corresponding *parameters*, or the percentage of the votes that each restaurant will receive from the population (i.e. all CU Boulder students). We will use confidence intervals to compute a range of values that reflects the uncertainty of our estimates.

The table `votes` contains the results of Oswaldo and Varun's survey.

```python
# Just run this cell
votes = pd.read_csv('data/votes.csv')
votes
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 6a (2 points)



Oswaldo and Varun notice that in their sample, Khow Thai has received more than half of the votes.    They are interested in obtaining a confidence interval for the percentage of votes that Khow Thai will receive from the population.


To begin, complete the function `one_resampled_percentage` below. It should return Khow Thai's ***percentage*** of votes after taking the original DataFrame as input and performing **one bootstrap sample** of it. Reminder that a percentage is between 0 and 100. 

*Note:* Input`tbl` will always be in the same format as the DataFrame `votes`.

*Hint:* Given a DataFrame of votes, how can you figure out what percentage of the votes are for a certain restaurant? **Be sure to use percentages, not proportions, for this question!**


```python
def one_resampled_percentage(tbl):
    ...

one_resampled_percentage(votes)
```

```python
grader.check("q6a")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 6b (2 points)

i).  Complete the `percentages_in_resamples(num_bootstraps)` function such that it simulates and returns a numpy array of length `num_bootstraps` elements, where each element represents a bootstrapped estimate of the percentage of voters who will vote for Khow Thai. You should use the `one_resampled_percentage` function you wrote above.


ii).  Run your function `percentages_in_resamples` for `num_bootstraps = 50000` and store the results in a numpy array called `resampled_percentages`.  Then create a density histogram of the entries in `resampled_percentages` array. **Set the bin number equal to 20 when plotting your histogram.**    Label both axes and include a title on your plot. 



```python
def percentages_in_resamples(num_bootstraps):
    
    ...

```

```python
...
# your code for histogram above this line
```

```python
grader.check("q6b")
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 6c (1 point)
Using the array `resampled_percentages`, find the values at the two edges of the middle 95% of the bootstrapped percentage estimates. Compute the lower and upper ends of the interval, named `khow_lower_bound` and `khow_upper_bound` respectively. 


```python
khow_lower_bound = ...
khow_upper_bound = ...
print(f"Bootstrapped 95% confidence interval for the percentage of Khow Thai voters in the population: [{khow_lower_bound:.2f}, {khow_upper_bound:.2f}]")
```

```python
grader.check("q6c")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 6d (1 point)
Suppose we produced 6,000 new samples (each one a new/distinct uniform random sample of 1,500 students) from the population and created a 95% confidence interval for the percentage of people who vote for Khow Thai from each one. Roughly how many of those 6,000 intervals do you expect will actually contain the true percentage of the population? 

Assign your answer to `true_percentage_intervals`.


```python
true_percentage_intervals = ...
```

```python
grader.check("q6d")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 6e (2 Points)
The staff also created 70%, 90%, and 99% confidence intervals from the same sample, but we forgot to label which confidence interval represented which percentages! **First**, match each confidence level (70%, 90%, 99%) with its corresponding interval in the cell below 

(e.g. __ % CI: [52.1, 54] $\rightarrow$ replace the blank with one of the three confidence levels).

**Then**, explain your thought process and how you came up with your answers. 

The intervals are below:

* [50.03, 55.94]
* [52.1, 54]
* [50.97, 54.99]



_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 8!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any websites you referenced*

### Submission Instructions

Before proceeding any further, **save this notebook.**

## IMPORTANT STEP BEFORE COMPILING THIS ASSIGNMENT. 

To make sure all of your plots render correctly in the PDF, please do the following steps now: 

1.) In the toolbar, go to Settings 

2.) Disable Autosave Documents 

**Double check that your PDF has all plots correctly showing before uploading to Gradescope - otherwise you WILL lose points**

```python
import simple_latex_checker as slc

nb = slc.Nb_checker()
nb.run_check("hw08.ipynb")
```

**TROUBLESHOOTING TIPS**  If you are having any issues compiling your assignment, please [**read through these troubleshooting tips first**](https://docs.google.com/document/d/1ndr3Wj1PSF5qzlLMaBJznwh6QGeEXjd5TAJ6nf9EJvo/edit?usp=sharing), then post any questions on Piazza. 

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 

* **Homework 8 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW 8 Autograded assignment in Gradescope.

* **Homework 8 Manually Graded**: Submit your hw08.PDF to the HW8 Manually Graded assignment in Gradescope.  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.  IF NOT, YOU WILL LOSE POINTS**   Also, **check** that all of your plots **and** all lines of your code are showing up in your PDF before submitting.  If not, you will not receive credit for your plots/code.  


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw08.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 1)
```
>>> assert get_hash(q1a_answer) == '6bb61e3b7bce0931da574d19d1d82c88'
```


Tests for q1b (points: 2)
```
>>> assert get_hash(round(q1b_answer, 3)) == '0a4a2fbfa050681b56c8bfce230ee97c'
```


Tests for q3a (points: 3)
```
>>> assert get_hash(round(sample_mean, 3)) == '7c4d074c8da37d2fbc276052bde68bee'
```
```
>>> assert get_hash(round(CI_lower, 3)) == '77807e922f3539eb5d2894d2db284d34'
```
```
>>> assert get_hash(round(CI_upper, 3)) == '9cf87aa3fe0926c8580b41dc953e18a7'
```


Tests for q3b (points: 1)
```
>>> assert get_hash(pop_avg_income) == '85c3f82be5c166a2b144a64af60c7a14'
```
```
>>> assert get_hash(pop_avg_in_CI) == '93cba07454f06a4a960172bbd6e2a435'
```


Tests for q3c (points: 4)
```
>>> assert get_hash(len(resampled_med)) == '14ee22eaba297944c96afdbe5b16c65b'
```
```
>>> import numpy as np
>>> np.random.seed(200)
>>> assert get_hash(round(one_resampled_med(sample), 0)) == 'bb8082cd8cd9c0df935ca3c352cbb843'
```
```
>>> import numpy as np
>>> np.random.seed(200)
>>> assert get_hash(round(np.array([one_resampled_med(sample) for i in range(100)]).mean(), 0))
```
```
>>> import numpy as np
>>> assert np.abs(np.percentile(resampled_med, 2.5) - 42833.5) < 90
```


Tests for q3d (points: 1)
```
>>> assert all([CI_med_lower == np.percentile(resampled_med, 2.5), CI_med_upper == np.percentile(resampled_med, 97.5)])
```


Tests for q3e (points: 1)
```
>>> assert get_hash(pop_median_income) == '5a7bd31d72318d2e45e508b561c453c5'
```
```
>>> assert get_hash(pop_median_in_CI) == '93cba07454f06a4a960172bbd6e2a435'
```


Tests for q5a (points: 3)
```
>>> assert get_hash(smallest) == 'b2f627fff19fda463cb386442eac2b3d'
```


Tests for q5b (points: 2)
```
>>> assert 1 <= sample_90 <= 3
```
```
>>> assert get_hash(sample_90) == 'c81e728d9d4c2f636f067f89cc14862c'
```


Tests for q5c (points: 2)
```
>>> assert get_hash(smallest_num) == 'e0c641195b27425bb056ac56f8953d24'
```


Tests for q5d (points: 1)
```
>>> assert np.isclose(shaded_area, 0.89899483)
```


Tests for q5e (points: [1, 1])
```
>>> assert get_hash(round(SE_99, 2)) == '312a77c20595d65b6cdb34b8da493ffb'
```
```
>>> assert get_hash(smallest_num_99) == 'd736bb10d83a904aefc1d6ce93dc54b8'
```


Tests for q6a (points: 2)
```
>>> assert type(one_resampled_percentage(votes)) in set([float, np.float64])
```
```
>>> assert one_resampled_percentage(votes) > 0
```
```
>>> assert 35 <= one_resampled_percentage(votes) <= 65
```
```
>>> np.random.seed(123)
>>> assert np.isclose(one_resampled_percentage(votes), 52.400000000000006)
```


Tests for q6b (points: [0, 0.5, 0.5])
```
>>> assert get_hash(len(resampled_percentages)) == '1017bfd4673955ffee4641ad3d481b1c'
```
```
>>> np.random.seed(123)
>>> assert np.isclose(percentages_in_resamples(2000).item(0), 52.400000000000006)
```
```
>>> np.random.seed(123)
>>> assert np.isclose(percentages_in_resamples(2000).item(10), 51.33333333333333)
```


Tests for q6c (points: 1)
```
>>> assert all([khow_lower_bound == np.percentile(resampled_percentages, 2.5), khow_upper_bound == np.percentile(resampled_percentages, 97.5)])
```


Tests for q6d (points: 1)
```
>>> assert get_hash(int(true_percentage_intervals)) == '84f5ddd735176becc72c3b1ff424149e'
```
